# Project Euler: Problem 52

https://projecteuler.net/problem=52

It can be seen that the number, $125874$, and its double, $251748$, contain exactly the same digits, but in a different order.

Find the smallest positive integer, $x$, such that $2x$, $3x$, $4x$, $5x$, and $6x$, contain the same digits.

First, write a function `same_digits(x,y)` that returns `True` if two integers `x` and `y` have the exact same set of digits and multiplicities and `False` if they have different digits.

<!-- grade_id: None, points: 0, grade: false, solution: true -->
```python
def same_digits(x, y):
    """Do the integers x and y have the same digits, regardless of order."""
    listx = list(str(x))
    listy = list(str(y))
    
    listx = sorted(listx)
    listy = sorted(listy)
    
    if listx == listy:
        return True
    else:
        return False
    
    
```

<!-- grade_id: projecteuler52a, points: 4, grade: true, solution: false -->
```python
assert same_digits('132', '321')
assert not same_digits('123', '3')
assert not same_digits('456', '0987654321')
```

Now use the `same_digits` function to solve this Euler problem. As you work on this problem, be careful to debug and test your code on small integers before trying it on the full search.

```python
assert same_digits("11100", "10011")
assert not same_digits("12345", "123455")
```

<!-- grade_id: None, points: 0, grade: false, solution: true -->
```python
for n in range(1, 1000000):
    if same_digits(n, 2*n) == True:
        if same_digits(n, 3*n) == True:
            if same_digits(n, 4*n) == True:
                if same_digits(n, 5*n) == True:
                    if same_digits(n, 6*n) == True:
                        print(n)
                        break
                        
```

Output:
```
142857
```

```python
print(142857, 2*142857, 3*142857, 4*142857, 5*142857, 6*142857)
```

Output:
```
142857 285714 428571 571428 714285 857142
```

<!-- grade_id: projecteuler52b, points: 6, grade: true, solution: false -->
```python
assert True # leave this cell to grade the solution
```