# Integration Exercise 3

## Imports

```python
%matplotlib inline
import matplotlib.pyplot as plt
import numpy as np
from scipy import integrate
```

## 2d polar integration

The 2d polar integral of a scalar function $f(r, \theta)$ is defined as:

$$ I(r_{max}) = \int_0^{r_{max}} \int_0^{2\pi} f(r, \theta) r dr d\theta $$

Write a function `integrate_polar(f, rmax)` that performs this integral numerically using `scipy.integrate.dblquad`.

<!-- grade_id: None, points: 0, grade: false, solution: true -->
```python
def integrate_polar(f, rmax):
    """Integrate the function f(r, theta) over r=[0,rmax], theta=[0,2*np.pi]"""
    #func = r * f(r)
    #not sure how to create the jacobian. Would require altering a function not yet defined?
    I, err = integrate.dblquad(func, 0, rmax, gfun = lambda x: 0, hfun = lambda x: 2*np.pi)
    return I
```

<!-- grade_id: integrationex04a, points: 10, grade: true, solution: false -->
```python
assert np.allclose(integrate_polar(lambda r,t: 1, 1.0), np.pi)
assert np.allclose(integrate_polar(lambda r, t: np.exp(-r)*(np.cos(t)**2), np.inf), np.pi)
```

Output:
```
<ipython-input-5-6f12576707cd> in <module>()
----> 1 assert np.allclose(integrate_polar(lambda r,t: 1, 1.0), np.pi)
      2 assert np.allclose(integrate_polar(lambda r, t: np.exp(-r)*(np.cos(t)**2), np.inf), np.pi)

<ipython-input-4-b71b7bb51f0d> in integrate_polar(f, rmax)
      1 def integrate_polar(f, rmax):
      2     """Integrate the function f(r, theta) over r=[0,rmax], theta=[0,2*np.pi]"""
----> 3     func = r * f(r)
      4     I, err = integrate.dblquad(func, 0, rmax, gfun = lambda x: 0, hfun = lambda x: 2*np.pi)
      5     return I

NameError: name 'r' is not defined
```

```python

```